Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Samuel Bourque"
COLLABORATORS = "..."
```

---

<img style='margin-right:0' src="http://dinfo.ca/logoDptInfo.jpg" width=300>

# Laboratoire #3
---

---
#### Question #1 (1 point)
On souhaite une fonction `determinerMotPlusLong(premierMot, secondMot)` qui déterminer, parmis deux mots fournis en paramètre, celui qui est le plus long des deux.  La fonction **retournera la chaîne de caractères** correspondant à ce mot.

Si les deux mots sont de même longueur, on retourne le premier mot.

<!-- grade_id: q1-code, points: 0, grade: false, solution: true -->
```python
def determinerMotPlusLong(premierMot, secondMot):
    if len(premierMot) > len(secondMot):
        retourMot = premierMot
    elif len(premierMot) == len(secondMot):
        retourMot = premierMot
    else:
        retourMot = secondMot
    return retourMot
```

<!-- grade_id: q1-tests, points: 1, grade: true, solution: false -->
```python
# Validations
assert determinerMotPlusLong('copain', 'ami') == 'copain'
assert determinerMotPlusLong('pair', 'impair') == 'impair'
assert determinerMotPlusLong('octet', 'linux') == 'octet'
assert determinerMotPlusLong('linux', 'octet') == 'linux'
```

---
#### Question #2 (1 point)

Le titre d'une personne (ou titre de civilité) permet d'identifier si on s'adresse à un homme ou une femme.  

Par exemple:

- Monsieur Kim Jong-un
- Madame Kim Kardashian

Définissez une fonction `determinerCodeSexe(titre)` qui permet de retourner "H" si le titre débute par "Monsieur" ou "F" si le titre débute par "Madame".

Lorsque le titre ne débute par ni l'un, ni l'autre, il faut retourner "?".

<!-- grade_id: q2-code, points: 0, grade: false, solution: true -->
```python
def determinerCodeSexe(titre):
    if titre.startswith("Monsieur"):
        codeSexe = 'H'
    elif titre.startswith("Madame"):
        codeSexe = 'F'
    else:
        codeSexe = '?'
    return codeSexe
```

<!-- grade_id: q2-tests, points: 1, grade: true, solution: false -->
```python
# Validations
assert determinerCodeSexe('Monsieur Kim Jong-un') == 'H'
assert determinerCodeSexe('Madame Kim Kardashian') == 'F'
assert determinerCodeSexe('Camarade Sasha Comanechi') == '?' 
```

---
#### Question #3 (1 point)

On souhaite disposer d'une fonction `compterVoyelles(mot)` qui retourne un nombre entier représentant le nombre de voyelles dans le mot.

Les voyelles sont 'a', 'e', 'i', 'o', 'u' et 'y'.  

> On ne considère pas les lettres accentuées.<br>
  Les lettres peuvent être en majuscule ou en minuscule

**Stratégie recommandée:**

- Compter le nombre de 'a' du mot et placer dans une variable
- Additionner à cette variable le nombre de 'e' du mot 
- Faire de même pour les autres voyelles


<!-- grade_id: q3-code, points: 0, grade: false, solution: true -->
```python
def compterVoyelles(mot):
    motMinuscule = mot.lower()
    compteur = motMinuscule.count('a')
    compteur += motMinuscule.count('e')
    compteur += motMinuscule.count('i')
    compteur += motMinuscule.count('o')
    compteur += motMinuscule.count('u')
    compteur += motMinuscule.count('y')
    return compteur
```

<!-- grade_id: q3-tests, points: 1, grade: true, solution: false -->
```python
# Validations
assert compterVoyelles('Yo-yo') == 4
assert compterVoyelles('Anticonstitutionnellement') == 10
assert compterVoyelles('aardvark') == 3
assert compterVoyelles('IRIS') == 2
assert compterVoyelles('galeries') == 4
assert compterVoyelles('ouistiti') == 5
assert compterVoyelles('oiseau') == 5
assert compterVoyelles('BRRR') == 0
```

---
#### Question #4 (1 point)

Déclarez une fonction `trouverLettreAlphabet(position)` qui retourne la lettre de l'alphabet (majuscule) qui correspond à la position demandée (entre 1 et 26).

> Considérez que la position qui sera fournie ne peut être en dehors de la plage (1 à 26).

**Stratégie recommandée:**

À l'intérieur de la fonction, composez une chaîne de caractères pour toutes les lettres de l'alphabet (dans l'ordre).  Ensuite, recherchez le caractère à la position demandée (attention car les indices débutent à zéro) dans cette grande chaîne.

<!-- grade_id: q4-code, points: 0, grade: false, solution: true -->
```python
def trouverLettreAlphabet(position):
    alpha = " ABCDEFGHIJKLMNOPQRSTUVWXYZ"   
    resultat = alpha [position]
    return resultat
```

```python
from ipywidgets import interact, interactive, fixed, interact_manual
import ipywidgets as widgets
interact(trouverLettreAlphabet, position=(1,26));
```

Output:
```
interactive(children=(IntSlider(value=13, description='position', max=26, min=1), Output()), _dom_classes=('wi…
```

<!-- grade_id: q4-tests, points: 1, grade: true, solution: false -->
```python
# Validations
assert trouverLettreAlphabet(2) == 'B'
assert trouverLettreAlphabet(5) == 'E'
assert trouverLettreAlphabet(25) == 'Y'
assert trouverLettreAlphabet(26) == 'Z'
assert trouverLettreAlphabet(1) == 'A'
assert trouverLettreAlphabet(3) == 'C'
assert trouverLettreAlphabet(4) == 'D'
assert trouverLettreAlphabet(6) == 'F'
assert trouverLettreAlphabet(7) == 'G'
assert trouverLettreAlphabet(8) == 'H'
assert trouverLettreAlphabet(9) == 'I'
assert trouverLettreAlphabet(10) == 'J'
assert trouverLettreAlphabet(11) == 'K'
assert trouverLettreAlphabet(12) == 'L'
assert trouverLettreAlphabet(13) == 'M'
assert trouverLettreAlphabet(14) == 'N'
assert trouverLettreAlphabet(15) == 'O'
assert trouverLettreAlphabet(16) == 'P'
assert trouverLettreAlphabet(17) == 'Q'
assert trouverLettreAlphabet(18) == 'R'
assert trouverLettreAlphabet(19) == 'S'
assert trouverLettreAlphabet(20) == 'T'
assert trouverLettreAlphabet(21) == 'U'
assert trouverLettreAlphabet(22) == 'V'
assert trouverLettreAlphabet(23) == 'W'
assert trouverLettreAlphabet(24) == 'X'
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-9-201265726500> in <module>()
      1 # Validations
      2 assert trouverLettreAlphabet(2) == 'B'
----> 3 assert trouverLettreAlphabet(5) == 'E'
      4 assert trouverLettreAlphabet(25) == 'Y'
      5 assert trouverLettreAlphabet(26) == 'Z'

AssertionError:
```

---
#### Question #5 (1 point)

Déclarez une fonction `trouverPositionAlphabet(lettre)` qui retourne la position dans l'alphabet de la lettre donnée en paramètre.

> Considérez que la lettre fournie peut être en minuscule ou en majuscule.<br>La lettre est obligatoirement une des 26 lettres de l'alphabet.

**Stratégie recommandée:**

À l'intérieur de la fonction, composez une chaîne de caractères pour toutes les lettres de l'alphabet (dans l'ordre).  Ensuite, recherchez la position dans la chaîne.  ATTENTION car la première lettre ('A') doit retourner la position 1 (et non pas la position 0).

<!-- grade_id: q5-code, points: 0, grade: false, solution: true -->
```python
def trouverPositionAlphabet(lettre):
    alpha = " ABCDEFGHIJKLMNOPQRSTUVWXYZ"   
    resultat = alpha.find(lettre.upper())
    return resultat
```

<!-- grade_id: q5-tests, points: 1, grade: true, solution: false -->
```python
# Validations
assert trouverPositionAlphabet('B') == 2
assert trouverPositionAlphabet('E') == 5
assert trouverPositionAlphabet('Y') == 25
assert trouverPositionAlphabet('Z') == 26
assert trouverPositionAlphabet('A') == 1
assert trouverPositionAlphabet('C') == 3
assert trouverPositionAlphabet('Z') == 26
assert trouverPositionAlphabet('a') == 1
assert trouverPositionAlphabet('c') == 3
assert trouverPositionAlphabet('x') == 24
```

---
#### Question #6 (1 point)

Déclarez une fonction `masquerVoyelles(mot)` qui retourne une chaîne où les voyelles ('a','e','i','o','u' et 'y') sont remplacées par un souligné.

Exemple:

> "voyelles" --> "v___ll_s"

*Considérez* que le mot reçu est toujours en minuscules.  

On ne considère que les voyelles sans accent.

**Stratégie recommandée:**

* Remplacez tous les 'a' du mot et placer le résultat dans une nouvelle variable
* Remplacez ensuite les 'e' présents dans cette nouvelle variable; placez le résultat à nouveau dans cette variable (remplace l'ancien contenu)
* Faites la même chose pour les autres voyelles


<!-- grade_id: q6-code, points: 0, grade: false, solution: true -->
```python
def masquerVoyelles(mot):
    reponse = mot.replace('a','_')
    reponse = reponse.replace('e','_')
    reponse = reponse.replace('i','_')
    reponse = reponse.replace('o','_')
    reponse = reponse.replace('u','_')
    reponse = reponse.replace('y','_')
    return reponse
```

<!-- grade_id: q6-tests, points: 1, grade: true, solution: false -->
```python
# Validations
assert masquerVoyelles('yo-yo') == '__-__'
assert masquerVoyelles('anticonstitutionnellement') == '_nt_c_nst_t_t__nn_ll_m_nt'
assert masquerVoyelles('aardvark') == '__rdv_rk'
assert masquerVoyelles('iris') == '_r_s'
assert masquerVoyelles('galeries') == 'g_l_r__s'
assert masquerVoyelles('ouistiti') == '___st_t_'
assert masquerVoyelles('oiseau') == '__s___'
assert masquerVoyelles('brrr') == 'brrr'
```

---
#### Question #7 (1 point)

La fonction `encadrer(mot)` permet d'**afficher** le mot entouré d'un cadre composée d'astéristiques.

Par exemple, `encadrer('informatique')` affichera:

```
****************
* informatique *
****************
```

**Stratégie possible:**

* Afficher un astéristique `n` fois, où `n` est la longueur de la chaîne + 4
* Afficher le mot, précédé d'un astérisque et d'un espace, puis suivi d'un espace et d'un astérisque
* Affichez la même ligne d'astérisque qu'au début

<!-- grade_id: q7-code, points: 1, grade: true, solution: true -->
```python
def encadrer(mot):
    nbAsterisque = len(mot) + 4
    print('*' * nbAsterisque)
    print("* %s *" % mot)
    print('*' * nbAsterisque)
```

```python
encadrer('informatique')
encadrer('anticonsitutionnellement')
encadrer('')
```

---
#### Question #8 (1 point)

Définissez une fonction `extraireExtension(nomFichier)` qui permet d'extraire l'extension et de la retourner (chaîne de caractères).

> Considérez que le nom du fichier est toujours valide et possède une extension.

Les extensions peuvent être de longueurs différentes.  La stratégie à appliquer est de rechercher à partir de la droite pour trouver la position du point.  Ensuite, faire un découpage de la chaîne après la position du point trouvé jusqu'à la fin de la chaîne.

<!-- grade_id: q8-code, points: 0, grade: false, solution: true -->
```python
def extraireExtension(nomFichier):
    chaine = nomFichier[nomFichier.find('.') + 1:]
    return chaine
```

<!-- grade_id: q8-tests, points: 1, grade: true, solution: false -->
```python
# Validations
assert extraireExtension('monalisa.gif') == 'gif'
assert extraireExtension('cv.docx') == 'docx'
assert extraireExtension('logo.png') == 'png'
assert extraireExtension('todo.txt') == 'txt'
assert extraireExtension('planCours.md') == 'md'
assert extraireExtension('programme.py') == 'py'
assert extraireExtension('laboratoire.ipynb') == 'ipynb'
```

---
#### Question #9 (1 point)

Définissez une fonction `extraireNomServeur(url)` qui permet d'extraire le nom d'un serveur (hôte) et de le retourner (chaîne de caractères).

> - Considérez que l'URL représente une adresse Web valide (protocole://hote/chemin).
- Le protocole est de longueur variable; il est délimité par "://"
- Considérez que les séparateurs sont toujours des barres obliques (slash '/') inclinées vers la droite; jamais des barres obliques inversées (backslash '\')
- Il y aura toujours une barre oblique (slash '/') après le nom de l'hôte

**Stratégie:**

* Trouver la position du deux-points (:)
* Extraire la sous-chaîne débutant deux positions plus loin que le deux-points (:) jusqu'à la fin de la chaîne; placez ce résultat dans une variable
* Extraire la sous-chaîne (de la variable) du début jusqu'à la première barre oblique (slash '/')

<!-- grade_id: q9-code, points: 0, grade: false, solution: true -->
```python
def extraireNomServeur(url):
    lien = url[url.find(":") + 3:]
    url = lien
    lien = url[:url.find("/")]
    return lien
```

<!-- grade_id: q9-tests, points: 1, grade: true, solution: false -->
```python
# Validations
assert extraireNomServeur('http://www.cegepthetford.ca/') == 'www.cegepthetford.ca'
assert extraireNomServeur('file://area51.mil/ultra/secret.txt') == 'area51.mil'
assert extraireNomServeur('https://c2z6.dinfo.cegepthetford.ca/examen1/e1.ipynb') == 'c2z6.dinfo.cegepthetford.ca'
```

---
#### Question #10 (1 point)

Définissez une fonction `extraireProtocole(uri)` qui permet d'extraire le protocole URI (scheme) et de le retourner (chaîne de caractères).

> - Considérez que l'URI est correctement formée (débute par `protocole:___`)
- Le protocole est de longueur variable; il est délimité par ":"

**Stratégie:**

* Extraire la chaîne jusqu'à la première occurrence du deux-point

<!-- grade_id: q10-code, points: 0, grade: false, solution: true -->
```python
def extraireProtocole(uri):
    lien = uri[:uri.find(':')]
    return lien
```

<!-- grade_id: q10-tests, points: 1, grade: true, solution: false -->
```python
# Validations
assert extraireProtocole('http://www.cegepthetford.ca/') == 'http'
assert extraireProtocole('file://area51.mil/ultra/secret.txt') == 'file'
assert extraireProtocole('https://c2z6.dinfo.cegepthetford.ca/examen1/e1.ipynb') == 'https'
assert extraireProtocole('tel://418-338-8591') == 'tel'
assert extraireProtocole('mailto:mguay@cegepthetford.ca') == 'mailto'
```

```python

```