Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Samuel Bourque"
COLLABORATORS = "..."
```

---

###### <img style='margin-right:0' src="http://dinfo.ca/logoDptInfo.jpg" width=300>

# Laboratoire #8
---

---
#### Question #1 (2 points)
On veut fusionner deux listes sans doublon à l'aide d'une nouvelle fonction `fusionnerListes(liste1, liste2)`.

Considérez qu'il n'y a aucune valeur en double dans chacun des listes fournies.

On souhaite retourner une nouvelle liste (fusion des valeurs, sans les doublons).

> Les deux listes en paramètre ne sont pas modifiées par cette fonction.

Stratégie:
- Cloner la première liste; cette nouvelle liste sera placée dans une variable `valeursFusion`.
- Pour chaque élément de la 2e liste, 
    - on valide si l'élément n'est pas présent dans la première liste (pour faire la négation d'une condition, utiliser l'opérateur `not`
    - si c'est le cas, on ajouter cet élément visité dans `valeursFusion`

Ne pas oublier de retourner `valeursFusion` à la fin de la fonction.

```python
# Exemple d'utilisation du `not`
chaine = '123abc'
if not chaine.isdigit():
    print('ne contient pas uniquement des chiffres')
```

<!-- grade_id: q1-code, points: 0, grade: false, solution: true -->
```python
def fusionnerListes(liste1, liste2):
    if not 
valeursFusion = liste1
raise NotImplementedError
```

<!-- grade_id: q1-tests, points: 2, grade: true, solution: false -->
```python
# Validations
assert 'fusionnerListes' in globals(), 'Aucune fonction ne s\'appelle `fusionnerListes`'
import inspect
assert len(inspect.signature(fusionnerListes).parameters) == 2, 'La fonction doit avoir deux paramètres'
assert list(inspect.signature(fusionnerListes).parameters.keys())[0] == 'liste1', 'Le 1er paramètre doit porter le nom `liste1`'
assert list(inspect.signature(fusionnerListes).parameters.keys())[1] == 'liste2', 'Le 2e paramètre doit porter le nom `liste2`'
assert fusionnerListes([1,2,3], [4,5,6]) == [1,2,3,4,5,6], 'Fusion #1 en échec'
assert fusionnerListes([1,2,3], [1,2,3]) == [1,2,3], 'Fusion #2 en échec'
assert fusionnerListes([1,2,3], [0,2,4]) == [1,2,3,0,4], 'Fusion #3 en échec'
assert fusionnerListes([1,2,3], []) == [1,2,3], 'Fusion #4 en échec'
assert fusionnerListes([], [4,5,6]) == [4,5,6], 'Fusion #5 en échec'
```

---
#### Question #2 (2 points)

Créer une fonction `calculerMoyenne(notes)` qui retourne la moyenne des notes.

Stratégie:
- définir une variable local `somme` initialisée à zéro
- additionner chaque valeur de la collection `notes`
- après le boucle, faites la division avec le nombre de notes dans la collection

> Considérez que la collection de notes ne sera jamais vide; il y aura au moins une note.

<!-- grade_id: q2-code, points: 0, grade: false, solution: true -->
```python
def calculerMoyenne(notes):
# your code here
raise NotImplementedError
```

<!-- grade_id: q2-tests, points: 2, grade: true, solution: false -->
```python
# Validations
assert calculerMoyenne( (10,10,10,10) ) == 10, 'Scores parfaits'
assert calculerMoyenne( (10,6,8,10) ) == 8.5, 'Moyenne normale'
assert calculerMoyenne( [5] ) == 5, 'Une seule note'
assert calculerMoyenne( (1,2,3,4,5,6,7,8,9,10) ) == 5.5, 'Moyenne atypique'
```

---
#### Question #3 (2 points)

On veut une fonction `compterValeursInf(valeurTemoin, valeurs)` qui compte le nombre d'éléments de la collection qui sont **sous la valeur témoin**.

Considérez:
- la collection de valeurs peut être vide
- il peut n'y avoir aucune valeur de la collection sous la valeur témoin

Stratégie:
- initialiser le compteur à zéro
- chaque valeur visitée sous la valeur témoin incrémente le compteur
- on retourne le compteur à la fin de la fonction

<!-- grade_id: cell-7e833c5c32d21094, points: 0, grade: false, solution: true -->
```python
def compterValeursInf(valeurTemoin, valeurs):
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-a244182cf521b7c4, points: 2, grade: true, solution: false -->
```python
assert compterValeursInf( 6, [10.0, 5.5, 2.0, 9.5, 8.0]) == 2, 'Test #1'
assert compterValeursInf( 2, [10.0, 5.5, 2.0, 9.5, 8.0]) == 0, 'Test #2'
assert compterValeursInf( 3, [10.0, 5.5, 2.0, 9.5, 8.0]) == 1, 'Test #3'
assert compterValeursInf(10, [10.0, 5.5, 2.0, 9.5, 8.0]) == 4, 'Test #4'
assert compterValeursInf(99, [10.0, 5.5, 2.0, 9.5, 8.0]) == 5, 'Test #5'
assert compterValeursInf(99, []) == 0, 'Test #6'
```

---
#### Question #4 (2 points)

On sait qu'une chaîne peut être décomposée en une collection de lettres avec `list(mot)`.

Définir la fonction `decomposerMot(mot)` qui retourne une collection de lettres en ordre alphabétique.

Créez une fonction qui retourne une collection de lettres d'un mot, mais ces lettres devront être toutes en minuscule et en ordre alphabétique.

<!-- grade_id: q4-code, points: 0, grade: false, solution: true -->
```python
def decomposerMot(mot):
# your code here
raise NotImplementedError
```

<!-- grade_id: q4-tests, points: 2, grade: true, solution: false -->
```python
# Validations
assert decomposerMot('abc') == ['a','b','c'], 'Test #1'
assert decomposerMot('') == [], 'Test #2'
assert decomposerMot('ABC') == ['a','b','c'], 'Test #3'
assert decomposerMot('lune') == ['e','l','n', 'u'], 'Test #4'
assert decomposerMot('Marco') == ['a','c','m', 'o', 'r'], 'Test #5'
assert decomposerMot('BIBI') == ['b','b','i', 'i'], 'Test #6'
```

---
#### Question #5 (2 points)

On veut une méthode `genererLotsCultivables(nbGraines)` qui génère une collection d'entiers représentant le nombre de graines par lot. 

Un lot est limité à 10 graines.

Si on génère les lots pour 35 graines, on obtiendra:
    `[10, 10, 10, 5]`

C'est-à-dire trois (3) lots complets et un lot partiel.

Stratégie:
- Aucune répétition n'est requise pour cette fonction
- À l'aide d'une division entière, compter le nombre de lots complets (`nbLotsComplets`)
- Le modulo permettra de compter les graines restantes (`nbGrainesRestantes`)
- Créer la collection des lots en dupliquant une collection `[10]` du nombre de lots complets
- Seulement s'il reste des graines, on ajoute ce chiffre aux lots
- Retourner les lots

<!-- grade_id: q5-code, points: 0, grade: false, solution: true -->
```python
def genererLotsCultivables(nbGraines):
# your code here
raise NotImplementedError
```

<!-- grade_id: q5-tests, points: 2, grade: true, solution: false -->
```python
# Validations
assert genererLotsCultivables(35) == [10, 10, 10, 5], 'Test #1'
assert genererLotsCultivables(18) == [10, 8], 'Test #2'
assert genererLotsCultivables(2)  == [2], 'Test #3'
assert genererLotsCultivables(50) == [10, 10, 10, 10, 10], 'Test #4'
```