Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

###### <img style='margin-right:0' src="http://dinfo.ca/logoDptInfo.jpg" width=300>

# Laboratoire #12 (Python)
---

---
#### Question #1 (1 point) -- fonction Lambda et filter

On veut une fonction `compterValeursInfFilter(valeurTemoin, valeurs)` qui compte le nombre d'éléments de la collection qui sont **sous la valeur témoin**.

Considérez:
- la collection de valeurs peut être vide
- il peut n'y avoir aucune valeur de la collection sous la valeur témoin

Stratégie:
- employer une fonction lambda avec la fonction filter afin d'obtenir une collection des éléments inférieurs à la valeur témoin
- convertir le résultat de `filter` en collection ordonnée (list)
- retourner le nombre d'élément de cette liste


<!-- grade_id: q1-reponse, points: 0, grade: false, solution: true -->
```python
def compterValeursInfFilter(valeurTemoin, valeurs):
    valeurs = filter ( lambda v:v%2==1, [])
    return list(valeurs) 
```

<!-- grade_id: q1-validation, points: 1, grade: true, solution: false -->
```python
assert compterValeursInfFilter( 6, [10.0, 5.5, 2.0, 9.5, 8.0]) == 2, 'Test #1'
assert compterValeursInfFilter( 2, [10.0, 5.5, 2.0, 9.5, 8.0]) == 0, 'Test #2'
assert compterValeursInfFilter( 3, [10.0, 5.5, 2.0, 9.5, 8.0]) == 1, 'Test #3'
assert compterValeursInfFilter(10, [10.0, 5.5, 2.0, 9.5, 8.0]) == 4, 'Test #4'
assert compterValeursInfFilter(99, [10.0, 5.5, 2.0, 9.5, 8.0]) == 5, 'Test #5'
assert compterValeursInfFilter(99, []) == 0, 'Test #6'
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-16-8ba0dd5421a9> in <module>()
----> 1 assert compterValeursInfFilter( 6, [10.0, 5.5, 2.0, 9.5, 8.0]) == 2, 'Test #1'
      2 assert compterValeursInfFilter( 2, [10.0, 5.5, 2.0, 9.5, 8.0]) == 0, 'Test #2'
      3 assert compterValeursInfFilter( 3, [10.0, 5.5, 2.0, 9.5, 8.0]) == 1, 'Test #3'
      4 assert compterValeursInfFilter(10, [10.0, 5.5, 2.0, 9.5, 8.0]) == 4, 'Test #4'
      5 assert compterValeursInfFilter(99, [10.0, 5.5, 2.0, 9.5, 8.0]) == 5, 'Test #5'

AssertionError: Test #1
```

---
#### Question #2 (1 point) -- liste en compréhensions

On veut une fonction `compterValeursInfCL(valeurTemoin, valeurs)` qui compte le nombre d'éléments de la collection qui sont **sous la valeur témoin**.

Considérez:
- la collection de valeurs peut être vide
- il peut n'y avoir aucune valeur de la collection sous la valeur témoin

Stratégie:
- employer une liste en compréhensions pour former une collection des valeurs inférieures à la valeur témoin
- retourner le nombre d'élément de cette liste

<!-- grade_id: q2-reponse, points: 0, grade: false, solution: true -->
```python
def compterValeursInfCL(valeurTemoin, valeurs):
# your code here
raise NotImplementedError
```

<!-- grade_id: q2-validation, points: 1, grade: true, solution: false -->
```python
assert compterValeursInfCL( 6, [10.0, 5.5, 2.0, 9.5, 8.0]) == 2, 'Test #1'
assert compterValeursInfCL( 2, [10.0, 5.5, 2.0, 9.5, 8.0]) == 0, 'Test #2'
assert compterValeursInfCL( 3, [10.0, 5.5, 2.0, 9.5, 8.0]) == 1, 'Test #3'
assert compterValeursInfCL(10, [10.0, 5.5, 2.0, 9.5, 8.0]) == 4, 'Test #4'
assert compterValeursInfCL(99, [10.0, 5.5, 2.0, 9.5, 8.0]) == 5, 'Test #5'
assert compterValeursInfCL(99, []) == 0, 'Test #6'
```

---
#### Question #3 (1 point) -- lambda et fonction `map`
Pour une collection ordonnée nommée `tableauPeriodique` (structure différente que celle du labo 11) dont chaque item est une collection d'association (clés: symbole, nom, noAtomique et masseAtomique) sur l'atome, vous devez extraire les noms des atomes (collection ordonnée de chaînes de caractère).

| nom       | symbole | noAtomique | masseAtomique |
|-----------|---------|------------|---------------|
| Hydrogène | H       | 1          | 1.008         |
| Hélium    | He      | 2          | 4.003         |
| Carbone   | C       | 6          | 12.000        |
| Azote     | N       | 7          | 14.000        |
| Oxygène   | O       | 8          | 16.000        |

Vous devez utiliser la fonction `map` et placer le résultat dans la variable `nomsAtomes`.

<!-- grade_id: q3-reponse, points: 0, grade: false, solution: true -->
```python
tableauPeriodique = [
    {
        'nom' : 'Hydrogène',
        'symbole' : 'H',
        'noAtomique' : 1,
        'masseAtomique' : 1.008
    },
    {
        'nom' : 'Hélium',
        'symbole' : 'He',        
        'noAtomique' : 2,
        'masseAtomique' : 4.003
    },
    {
        'nom' : 'Carbone',
        'symbole' : 'C',        
        'noAtomique' : 6,
        'masseAtomique' : 12.000
    },    
    {
        'nom' : 'Azote',
        'symbole' : 'N',        
        'noAtomique' : 7,
        'masseAtomique' : 14.000
    },
    {
        'nom' : 'Oxygène',
        'symbole' : 'O',
        'noAtomique' : 8,
        'masseAtomique' : 16.000
    }
]

# your code here
raise NotImplementedError
```

<!-- grade_id: q3-validation, points: 1, grade: true, solution: false -->
```python
# Validations
assert 'nomsAtomes' in globals(), 'Aucune variable portant le nom `nomsAtomes`'
assert type(nomsAtomes) == list, 'Mauvais type pour la collection'
assert len(nomsAtomes) == 5, 'Il n\'y a pas cinq (5) atomes dans votre collection'
assert 'Hydrogène' in nomsAtomes, 'Pas d''hydrogène'
assert 'Hélium' in nomsAtomes, 'Pas d''hélium'
assert 'Carbone' in nomsAtomes, 'Pas de carbone'
assert 'Azote' in nomsAtomes, 'Pas d''azote'
assert 'Oxygène' in nomsAtomes, 'Pas d''oxygène'
```

---
#### Question #4 (1 point) -- liste en compréhensions
Pour une collection ordonnée nommée `tableauPeriodique` (structure différente que celle du labo 11) dont chaque item est une collection d'association (clés: symbole, nom, noAtomique et masseAtomique) sur l'atome, vous devez extraire les symboles des atomes (collection ordonnée de chaînes de caractère).

| nom       | symbole | noAtomique | masseAtomique |
|-----------|---------|------------|---------------|
| Hydrogène | H       | 1          | 1.008         |
| Hélium    | He      | 2          | 4.003         |
| Carbone   | C       | 6          | 12.000        |
| Azote     | N       | 7          | 14.000        |
| Oxygène   | O       | 8          | 16.000        |

Vous devez utiliser une liste en compréhension et placer le résultat dans la variable `symbolesAtomes`.

<!-- grade_id: q4-reponse, points: 0, grade: false, solution: true -->
```python
tableauPeriodique = [
    {
        'symbole' : 'H',
        'nom' : 'Hydrogène',
        'noAtomique' : 1,
        'masseAtomique' : 1.008
    },
    {
        'symbole' : 'He',        
        'nom' : 'Hélium',
        'noAtomique' : 2,
        'masseAtomique' : 4.003
    },
    {
        'symbole' : 'C',        
        'nom' : 'Carbone',
        'noAtomique' : 6,
        'masseAtomique' : 12.000
    },    
    {
        'symbole' : 'N',        
        'nom' : 'Azote',
        'noAtomique' : 7,
        'masseAtomique' : 14.000
    },
    {
        'symbole' : 'O',
        'nom' : 'Oxygène',
        'noAtomique' : 8,
        'masseAtomique' : 16.000
    }
]

# your code here
raise NotImplementedError
```

<!-- grade_id: q4-validation, points: 1, grade: true, solution: false -->
```python
# Validations
assert 'symbolesAtomes' in globals(), 'Aucune variable portant le nom `nomsAtomes`'
assert type(symbolesAtomes) == list, 'Mauvais type pour la collection'
assert len(symbolesAtomes) == 5, 'Il n\'y a pas cinq (5) atomes dans votre collection'
assert 'H' in symbolesAtomes, 'Pas d''hydrogène'
assert 'He' in symbolesAtomes, 'Pas d''hélium'
assert 'C' in symbolesAtomes, 'Pas de carbone'
assert 'N' in symbolesAtomes, 'Pas d''azote'
assert 'O' in symbolesAtomes, 'Pas d''oxygène'
```