Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```javascript
NAME = "Samuel Bourque"
COLLABORATORS = "..."
```

---

###### <img style='margin-right:0' src="http://dinfo.ca/logoDptInfo.jpg" width=300>

# Laboratoire #13 (JavaScript)
---

#### Question #1 (2 points)

Créez une collection d'association (Map) ayant comme clé le numéro d'employé et comme valeur associée un objet JavaScript.  

Cette collection doit être référencée par la variable `employes`.

Les clés de l'objet JavaScript doivent être un symbole (et non une chaîne).

L'objet JavaScript possède 3 clés: 
- nom
- nbHeures
- tauxHoraire

| noEmpl | nom             | nbHeures | tauxHoraire |
|--------|-----------------|----------|-------------|
| 11122  | 'Adam Bernard'  |   0.0    |  15.00      |
| 22233  | 'Charles Demers'|   15.5   |  20.00      |
| 33344  | 'Eric Fillion'  |   7.0    |  12.25      |


<!-- grade_id: q1-reponse, points: 1, grade: true, solution: true -->
```javascript
employes = new Map( [
    {noEmpl : 11122, nom : "Adam Bernard", nbHeures : 0.0, tauxHoraire : 15.00},
    {noEmpl : 22233, nom : "Charles Demers", nbHeures : 15.5, tauxHoraire : 20.00},
    {noEmpl : 33344, nom : "Eric Fillion", nbHeures : 7.0, tauxHoraire : 12.25} 
   ] );
```

Output:
```
Map { undefined => undefined }
```

<!-- grade_id: q1-validation, points: 1, grade: true, solution: false -->
```javascript
var assert = require('assert');
assert(typeof employes !== 'undefined', 'Pas déclaré la variable `employes`');
assert(employes instanceof Map, '`employes` doit être de type Map');
assert(employes.size == 3, 'Doit y avoir trois éléments dans la collection');
assert(employes.has(11122), 'Clé 11122 absente');
assert(employes.has(22233), 'Clé 22233 absente');
assert(employes.has(33344), 'Clé 33344 absente');
assert(employes.get(11122).hasOwnProperty('nom'));
assert(employes.get(22233).hasOwnProperty('nom'));
assert(employes.get(33344).hasOwnProperty('nom'));
assert(employes.get(11122).hasOwnProperty('nbHeures'));
assert(employes.get(22233).hasOwnProperty('nbHeures'));
assert(employes.get(33344).hasOwnProperty('nbHeures'));
assert(employes.get(11122).hasOwnProperty('tauxHoraire'));
assert(employes.get(22233).hasOwnProperty('tauxHoraire'));
assert(employes.get(33344).hasOwnProperty('tauxHoraire'));
assert(employes.get(11122).nom === 'Adam Bernard');
assert(employes.get(22233).nom === 'Charles Demers');
assert(employes.get(33344).nom === 'Eric Fillion');
```

Output:
```
at process.emit (events.js:191:7)
    at process.nextTick (internal/child_process.js:758:12)
    at _combinedTickCallback (internal/process/next_tick.js:73:7)
```

---
#### Question #2 (2 points) 

Pour la collection des employés, on veut calculer le salaire brut (nombre d'heures x taux horaire).

NOTE: On ne fait pas de calcul pour les heures supplémentaires.

Résultats attendus:
```
Adam Bernard (11122) -- salaire brut: 0.00$
Charles Demers (22233) -- salaire brut: 310.00$
Eric Fillion (33344) -- salaire brut: 85.75$
```

<!-- grade_id: q2-reponse, points: 2, grade: true, solution: true -->
```javascript
employes
" -- "
nbHeures * tauxHoraire
```

---
#### Question #3 (3 points)

Chargez le fichier JSON nommé 'films.json'; ce fichier se trouve dans le répertoire courant du laboratoire.

**Utilisez `JSONEditorOnline.org` pour comprendre la structure**

Affichez le titre de chaque film.

Résultats attendus (début de la liste):

```
Venom
Bohemian Rhapsody
Les Indestructibles 2
The Predator
A Star Is Born
```

<!-- grade_id: q3-reponse, points: 3, grade: true, solution: true -->
```javascript
var fs = require('fs');
var chaineJSON = fs.readFileSync('films.json', 'utf8');
// your code here
throw new Error();
```

---
#### Question #4 (3 points)

Quel est le film ayant eu le plus de votes (vote_count)?  
> Par programmation, faites la recherche dans les résultats JSON.

Assignez les valeurs aux variables `titre`, `description` et `poster` (url de l'image de l'affiche).  

Exemple de résultats attendus (<span style="background-color:yellow">note: pas le film recherché</span>):
<div style="border:1px solid black; padding:10px">
  <h1>Venom</h1>
  <p>Des symbiotes débarquent sur la Terre, parmi eux, Venom, qui va s'allier avec Eddie Brock. De son côté, un puissant scientifique tente de faire évoluer l'humanité avec les symbiotes, le duo d'anti-héros va devoir tout faire pour l'arrêter !<p/>
  <img src="https://image.tmdb.org/t/p/w185/vVusHIRlyyFVS42XnqZso2wGKr.jpg">
</div>

<!-- grade_id: q4-reponse, points: 1, grade: true, solution: true -->
```javascript
var poster = '/vVusHIRlyyFVS42XnqZso2wGKr.jpg';
var titre = 'titre';
var description = 'Lorem ipsum dolor sit amet...';
// your code here
throw new Error();
$$.html(`
  <h1>${titre}</h1>
  <p>${description}<p/>
  <img src="https://image.tmdb.org/t/p/w185/${poster}">`);
```

<!-- grade_id: cell-f85ee74b174cdd51, points: 2, grade: true, solution: false -->
```javascript
var assert = require('assert');
assert(poster.endsWith('Z96VLkqEZo3.jpg'), 'Mauvaise image pour l\'affiche');
assert(titre.startsWith('Les Ind'), 'Mauvais titre');
assert(description.startsWith('Notre famille de super-héros'), 'Mauvaise description')
```