Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Samuel Bourque"
COLLABORATORS = "..."
```

---

##### <img style='margin-right:0' src="http://dinfo.ca/logoDptInfo.jpg" width=300>

# Laboratoire #14 (Python)
---

#### Question #1 (2 points)

Le fichier `nombres.txt` est composée de nombres (un nombre par ligne).  

> Vous pouvez exécuter la commande `cat` de la cellule suivante pour consulter le contenu.

Il faut additionner tous ces nombres avec la variable `total` et afficher ce total.

**Résultat attendu**:
```
Total: 121
```

```python
!cat nombres.txt
```

Output:
```
30
40
12
34
5
```

<!-- grade_id: q1-reponse, points: 1, grade: true, solution: true -->
```python
with open ('nombres.txt','r') as f:
    compteur = 0
    total = f.readlines()
    for unTotal in total:
        result = f + total
    print("Total:", total)
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-52-ccafc14a6208> in <module>()
      3     total = f.readlines()
      4     for unTotal in total:
----> 5         result = f + total
      6     print("Total:", total)

TypeError: unsupported operand type(s) for +: '_io.TextIOWrapper' and 'list'
```

<!-- grade_id: q1-validation, points: 1, grade: true, solution: false -->
```python
assert total == 121, 'Probleme de calcul pour le total'
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-22-d72a1c2cb17d> in <module>()
----> 1 assert total == 121, 'Probleme de calcul pour le total'

AssertionError: Probleme de calcul pour le total
```

---
#### Question #2 (2 points) 

Vous devez compléter la fonction `fusionnerFichiers(nomFicA, nomFicB)` qui recopie les lignes du premier fichier dans `fusion.txt`, puis ajoute toutes les lignes du second fichier à la suite.

<!-- grade_id: q2-reponse, points: 1, grade: true, solution: true -->
```python
def fusionnerFichiers(nomFicA, nomFicB):
with open ('fusion.txt','') as f:
```

```python
fusionnerFichiers('premier.txt','second.txt')
!cat fusion.txt
```

<!-- grade_id: cell-d5523b2184cd5999, points: 1, grade: true, solution: false -->
```python
nbLignes = !wc -l fusion.txt
assert nbLignes[0].split()[0] == '6', 'Pas six lignes dans la fusion'
```

---
#### Question #3 (2 points)

Nous avons un fichier `notes.txt` qui comprend les résultats obtenus à cinq questions pour un exercice.

> Vous pouvez exécuter la cellule suivante pour voir la structure de ce fichier (les éléments sont séparés par des virgules).

La pondération pour les questions sont:
- Q1: 2 points
- Q2: 1 point
- Q3: 1 point
- Q4: 2 points
- Q5: 4 points

L'exercice est sur 10 points.

On veut afficher à l'écran la note obtenue pour chaque étudiant.

*Résultats attendus*:
```
Adam      : 10.0/10
Bernard   : 5.0/10
Charles   : 9.6/10
```

```python
# Structure du fichier
!cat notes.txt
```

Output:
```
Adam,2,1,1,2,4
Bernard,1.5,0,0.5,2,1
Charles,2,0.8,0.8,2,4
```

<!-- grade_id: q3-reponse, points: 2, grade: true, solution: true -->
```python
# your code here
raise NotImplementedError
# Format d'affichage: "%-10s: %.1f/10" 
```

---
#### Question #4 (2 points)

Dans le fichier `alphabet.txt`, nous avons tous les mots de l'alphabet international.

Vous devez compléter la fonction `chargerAlphabet()` qui retourne une collection d'association; la clé est la première lettre du mot.  On ne considère que les lettres de l'alphabet.

```python
!cat alphabet.txt
```

Output:
```
alpha
bravo
charlie
delta
echo
foxtrot
golf
hotel
india
juliet
kilo
lima
mike
november
oscar
papa
quebec
romeo
sierra
tango
uniform
victor
whiskey
x-ray
yankee
zulu
```

<!-- grade_id: q4-reponse, points: 1, grade: true, solution: true -->
```python
def chargerAlphabet():
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-7143739e3c20519b, points: 1, grade: true, solution: false -->
```python
validAlpha = chargerAlphabet()
assert len(validAlpha) == 26, 'Doit y avoir 26 lettres'
assert 'a' in validAlpha, 'Pas de `a` comme clé'
assert 'z' in validAlpha, 'Pas de `z` comme clé'
assert validAlpha['q'] == 'quebec', 'quebec'
assert validAlpha['a'] == 'alpha', 'alpha'
assert validAlpha['z'] == 'zulu', 'zulu'
```

---
#### Question #5 (2 points)

Le fichier `remorquages.csv` est très volumineux.  

**Votre solution doit obligatoirement** utiliser la méthode **readline()**.

> Il est <span style="color:red">interdit</span> de prendre `readlines()` ou `read()`.

Voici à quoi ressemble le début de ce fichier:

```
DATE_ORIGINE,LONGITUDE_ORIGINE,LATITUDE_ORIGINE,RUE_ORIGINE,SECTEUR_ORIGINE,ARRONDISSEMENT_ORIGINE,DATE_DESTINATION,LONGITUDE_DESTINATION,LATITUDE_DESTINATION,RUE_DESTINATION,SECTEUR_DESTINATION,ARRONDISSEMENT_DESTINATION,MOTIF_REMORQUAGE
"2015-12-31T00:23:00",-73.5872033593,45.5249805457,"rue Drolet ",,"Plateau-Mont-Royal","2015-12-31T00:25:00",-73.5867910237,45.5251307666,"rue de Grand-Pré ",,"Plateau-Mont-Royal","Constat d'infraction"
"2015-12-31T00:24:00",-73.6222559028,45.5857531685,"boulevard des Grandes-Prairies ",,"Saint-Léonard","2015-12-31T00:30:00",-73.6210045429,45.5869748675,"boulevard des Grandes-Prairies ",,"Saint-Léonard","Constat d'infraction"
"2015-12-31T00:25:00",-73.6006920362,45.535027244,"rue Saint-Hubert ",,"Rosemont - La Petite-Patrie","2015-12-31T00:32:00",-73.6008810993,45.5351626927,"rue Saint-Hubert ",,"Rosemont - La Petite-Patrie","Constat d'infraction"
```

La toute première ligne décrit les champs; il faut donc sauter par-dessus cette ligne avant de débuter le traitement.

Votre mandat consiste à compter le nombre d'infractions (nombre de lignes présentes dans ce fichier) en 2017.

**Résultat attendu**:
```
Nombre d'infractions en 2017: 53033
```

<!-- grade_id: q5-reponse, points: 1, grade: true, solution: true -->
```python
# your code here
raise NotImplementedError
print("Nombre d'infractions en 2017:", compteur)
```

<!-- grade_id: q5-validation, points: 1, grade: true, solution: false -->
```python
assert compteur == 53033
```