Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```javascript
NAME = "Samuel Bourque"
COLLABORATORS = "..."
```

---

###### <img style='margin-right:0' src="http://dinfo.ca/logoDptInfo.jpg" width=300>

# Laboratoire #12 (JavaScript)
---

#### Question #5 (1 point)

Vous devez afficher, en utilisant la méthode d'itération `forEach()`, chaque pays du G7 comme suit:

>    ________ est membre du G7

où le souligné est remplacé par le nom du pays    

<!-- grade_id: q5-reponse, points: 1, grade: true, solution: true -->
```javascript
var paysG7 = ['États-Unis', 'Japon', 'Allemagne', 'France', 'Royaume-Uni', 'Italie', 'Canada'];
paysG7.forEach( (unPaysG7 => {console.log(unPaysG7,"est membre du G7")}))
```

Output:
```
États-Unis est membre du G7
Japon est membre du G7
Allemagne est membre du G7
France est membre du G7
Royaume-Uni est membre du G7
Italie est membre du G7
Canada est membre du G7
```

---
#### Question #6 (1 point) -- méthode d'itération `map`

Créez une fonction `calculerTaxes(listePrix)` qui reçoit un collection de prix et retourne la taxe calculée pour chacun dans une nouvelle collection.  La collection d'origine n'est pas modifiée par votre méthode.

Le calcul doit être arrondi à 2 décimales.  **La collection retournée doit être composée de nombres et non pas de chaînes de caractères**

Stratégie:
* utiliser une méthode d'itération `map`
* pour chaque item visité
    * calculer la taxe, faire l'arrondi à 2 décimales et ajouter ce montant (**numérique**) à la nouvelle collection

<!-- grade_id: q6-reponse, points: 0, grade: false, solution: true -->
```javascript
function calculerTaxes(listePrix) {
    return listePrix.map((uneTaxe)=>parseFloat(uneTaxe*0.14975).toFixed(2));
}
```

```javascript
calculerTaxes([10.00, 100.00, 50.00, 40.00, 30.00, 20.00]);
// Attendu: [1.5, 14.97, 7.49, 5.99, 4.49, 3]
```

Output:
```
[ '1.50', '14.97', '7.49', '5.99', '4.49', '3.00' ]
```

<!-- grade_id: q6-tests, points: 1, grade: true, solution: false -->
```javascript
var assert = require('assert');
// Validations
assert( calculerTaxes([10.00, 100.00, 50.00, 40.00, 30.00, 20.00]).every( (v,i) => v == [ 1.5, 14.97, 7.49, 5.99, 4.49, 3 ][i]), 'Problème pour les valeurs attendues');
```

---
#### Question #7 (1 point) -- filter

Créez une fonction `denombrerElevesPresents(listePresences)` qui compte le nombre d'items qui sont "présent".

Retourne 0 pour une collection vide ou si tous les élèves sont absents!

Stratégie:
- à l'aide d'un filtre (méthode d'itération `filter`), extraire les items ayant la valeur `présent`
- retourner le nombre d'items extraits

<!-- grade_id: q7-reponse, points: 0, grade: false, solution: true -->
```javascript
function denombrerElevesPresents(listePresences) {
    compteur = 0;
    listePresences.filter( (mot) => mot.length > 6 );
    return compteur;
}
```

```javascript
console.log("Attendu: 3, réel: ", denombrerElevesPresents(["présent", "absent", "absent", "présent", "présent"]) );
console.log("Attendu: 5, réel: ", denombrerElevesPresents(["présent", "présent", "présent", "présent", "présent"]) );
console.log("Attendu: 0, réel: ", denombrerElevesPresents(["absent", "absent", "absent", "absent", "absent"]) );
console.log("Attendu: 0, réel: ", denombrerElevesPresents([]) );
```

Output:
```
Attendu: 3, réel:  0
Attendu: 5, réel:  0
Attendu: 0, réel:  0
Attendu: 0, réel:  0
```

<!-- grade_id: q7-tests, points: 1, grade: true, solution: false -->
```javascript
var assert = require('assert');
// Validations
assert(denombrerElevesPresents(["présent", "absent", "absent", "présent", "présent"]) == 3 );
assert(denombrerElevesPresents(["présent", "présent", "présent", "présent", "présent"]) == 5 );
assert(denombrerElevesPresents(["absent", "absent", "absent", "absent", "absent"]) == 0);
assert(denombrerElevesPresents([]) == 0 );
```

Output:
```
at process.emit (events.js:191:7)
    at process.nextTick (internal/child_process.js:758:12)
    at _combinedTickCallback (internal/process/next_tick.js:73:7)
```

---
#### Question #8 (1 point) -- every

Créez une fonction `validerTousPresents(listePresences)` qui retourne vrai si tous les élèves sont présents (aucune absence).

Stratégie:
- à l'aide de la méthode d'itération `every`, vérifier si l'état est `présent`
- retourner le résultat de cette méthode

<!-- grade_id: q8-reponse, points: 0, grade: false, solution: true -->
```javascript
function validerTousPresents(listePresences) {
    return listePresences.every("présent");
}
```

<!-- grade_id: q8-validation, points: 1, grade: true, solution: false -->
```javascript
var assert = require('assert');
// Validations
assert(!validerTousPresents(["présent", "absent", "absent", "présent", "présent"]), 'Trois élèves présents sur cinq');
assert(validerTousPresents(["présent", "présent", "présent", "présent", "présent"]), 'Tous les élèves sont présents');
assert(!validerTousPresents(["absent", "absent", "absent", "absent", "absent"]), 'Aucun élève présent si tous absents');
```

Output:
```
at onMessage ([eval]:789:13)
    at emitTwo (events.js:106:13)
    at process.emit (events.js:191:7)
```

---
#### Question #9 (1 point)

On veut une fonction `compterValeursInf(valeurTemoin, valeurs)` qui compte le nombre d'éléments de la collection qui sont **sous la valeur témoin**.

Considérez:
- la collection de valeurs peut être vide
- il peut n'y avoir aucune valeur de la collection sous la valeur témoin

Stratégie:
- utiliser une méthode `filter` pour extraire la liste des valeurs inférieures à la valeur témoin
- retourner la longueur de cette liste extraite

<!-- grade_id: q9-reponse, points: 0, grade: false, solution: true -->
```javascript
function compterValeursInf(valeurTemoin, valeurs) {
    valeurTémoin.filter(valeurTémoin > valeurs);
    return len(valeurTémoin < valeur)
}
```

<!-- grade_id: q9-validation, points: 1, grade: true, solution: false -->
```javascript
var assert = require('assert');
assert(compterValeursInf( 6, [10.0, 5.5, 2.0, 9.5, 8.0]) === 2, 'Test #1');
assert(compterValeursInf( 2, [10.0, 5.5, 2.0, 9.5, 8.0]) === 0, 'Test #2');
assert(compterValeursInf( 3, [10.0, 5.5, 2.0, 9.5, 8.0]) === 1, 'Test #3');
assert(compterValeursInf(10, [10.0, 5.5, 2.0, 9.5, 8.0]) === 4, 'Test #4');
assert(compterValeursInf(99, [10.0, 5.5, 2.0, 9.5, 8.0]) === 5, 'Test #5');
assert(compterValeursInf(99, []) === 0, 'Test #6');
```

Output:
```
at emitTwo (events.js:106:13)
    at process.emit (events.js:191:7)
    at process.nextTick (internal/child_process.js:758:12)
```

---
#### Question #10 (1 point)

Pour une collection ordonnée nommée `tableauPeriodique` dont chaque item est une collection ordonnée (indice 0: nom, indice 1 : symbole, indice 2: noAtomique et indice: 3 masseAtomique), vous devez rechercher l'array décrivant l'atome pour le nom 'Carbone' à l'aide de la méthode `find`.  Placez le résultat dans une variable nommée `atomeCarbone`.

| nom       | symbole | noAtomique | masseAtomique |
|-----------|---------|------------|---------------|
| Hydrogène | H       | 1          | 1.008         |
| Hélium    | He      | 2          | 4.003         |
| Carbone   | C       | 6          | 12.000        |
| Azote     | N       | 7          | 14.000        |
| Oxygène   | O       | 8          | 16.000        |

<!-- grade_id: q10-reponse, points: 0, grade: false, solution: true -->
```javascript
tableauPeriodique = [
    ['Hydrogène', 'H',  1, 1.008],
    ['Hélium',    'He', 2, 4.003],
    ['Carbone',   'C',  6, 12.000],
    ['Azote',     'N',  7, 14.000],
    ['Oxygène',   'O',  8, 16.000]
]

var atomeCarbone = 
```

```javascript
var fluctuations = [50, 60, 55, 58, 52, 70, 80, 48];
function estFluctuation(mesure) {
    return mesure < 50 || mesure > 60;
}
console.log('Indice première fluctuation:', fluctuations.findIndex(estFluctuation));
```

<!-- grade_id: q10-validation, points: 1, grade: true, solution: false -->
```javascript
var assert = require('assert');
assert(atomeCarbone != undefined, "Pas trouvé d'atome: undefined");
assert(atomeCarbone[0] === 'Carbone', "Mauvais nom pour l'atome");
assert(atomeCarbone[1] === 'C', "Mauvais symbole");
assert(atomeCarbone[2] === 6, "Mauvais numéro atomique");
assert(atomeCarbone[3] === 12.0, "Mauvaise masse atomique");
```